# Dados do aluno
Preencha seus dados nas variáveis da célula abaixo para continuar.

```python
# Preencha seu nome e código do aluno aqui
nome_aluno   = ""   # Substitua aqui pelo seu nome, exemplo:   nome_aluno = "Edson Crosué das Dores"
codigo_aluno = None # Substitua aqui pelo seu código, exemplo: codigo_aluno = 1234567
```

<!-- grade_id: cell-a3a35353a136090d, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert nome_aluno != "" # Você precisa preencher seu nome no formulário
assert codigo_aluno is not None # Você precisa preencher seu código de aluno
```

# Sobre
Python é uma linguagem de programação que busca aumentar o nível de abstração da implementação.

Assim como em outras linguagens, o conceito de funções, estruturas de decisão e de repetição em Python serão usado constantemente e, nesta disciplina, isso não será diferente. Além deles, tipos iteráveis, como listas e tuplas e dicionários são usados quase que em todo script Python.

Desta forma, neste notebook o objetivo é guiar o estudo sobre o básico de implementação em Python.

# Exercício 1
Implemente uma função chamada `eh_positivo`, a qual retorna `True` caso a variável de entrada seja maior do que zero e retorna `False` caso contrário.

**Lembrete:** a sintaxe da declaração de uma função é:
```python
def nome_funcao(variavel_1, variavel_2, ..., variavel_N):
    # Código
    ...
    # Fim código
    return resultado # Retorna algum valor calculado, caso haja algo para retornar
```

**Lembrete:** a sintaxe da estrutura de decisão em Python é:
```python
if condicao_1:
    bloco_condicao_1
elif condicao_2: # o 'elif' é opcional
    bloco_condicao_2
...
else: # Caso nenhuma das condições seja verdadeira, cai no 'else'
    bloco_else
```

<!-- grade_id: cell-7e915fa7896192dd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Use esta célula para testar sua função. Altere a variável 'val_teste' conforme queira
val_teste = 10
print(eh_positivo(val_teste))
```

<!-- grade_id: cell-d664ce978f6e299e, points: 2, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert eh_positivo(1) is not None # Feedback: por um acaso você se lembrou de colocar 'return'?
assert eh_positivo(5) == True # Feedback: retornou que 5 não é positivo. Confira seu código.
assert eh_positivo(0) == False # Feedback: retornou que 0 é positivo. Confira seu código.
assert eh_positivo(-1) == False # Feedback: retornou que -1 é positivo. Confira seu código.

```

# Exercício 2
A função implementada resolve o problema caso o valor fornecido seja um número. Se quiséssemos avaliar cada elemento de uma lista seria necessário implementar uma estrutura de repetição por fora da função e isso poderia deixar o código que chamou a função um pouco menos legível.

Desta forma, **implemente uma função nomeada `eh_positivo_lista` em que a variável recebida é uma lista e que fornece como saída outra lista que tem o resultado da função `eh_positivo` aplicada a cada elemento da lista**. 

Por exemplo, se fornecermos a lista `[10, 0, -5]` como entrada, a saída esperada é a lista `[True, False, False]`.

**Dica:** Comece criando uma lista vazia e concatene a ela o resultado de cada elemento da lista de entrada. Não se esqueça de retornar o resultado usando `return`!

**Lembrete:** Para criar uma lista, faça: `lista = []` e para concatenar elementos a ela use `lista.append(valor)`, em que `valor` é uma variável ou um literal.

**Lembrete:** Existem duas formas de fazer repetição em Python: usando `for` e `while`: 
```python
while condicao:
    bloco_while
```
```python
for valor in iteravel:
    bloco_for
```
No caso do `while`, para acessar cada elemento da lista você precisará de uma variável auxiliar (iniciada em 0) e que irá ser incrementada até o número de elementos -1.

No caso do `for`, a cada iteração cada elemento da variável `iteravel` é acessado em sequência e atribuído à variável `valor`.

<!-- grade_id: cell-1e6b56b8662565ed, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Use esta célula para testar sua função. Altere a variável 'val_teste' conforme queira
val_teste = [10,0,-5]
print(eh_positivo_lista(val_teste))
```

<!-- grade_id: cell-de40a2aada928521, points: 2, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
resultado = eh_positivo_lista([10,0,-5])
gabarito = [True, False, False]
assert all([ gabarito[i] == resultado[i] for i in range(3)  ]) # Feedback: Algum dos elementos da lista de teste teve resultado diferente. Confira seu código.

resultado = eh_positivo_lista([50,-7,10,5])
gabarito = [True, False, True, True]
assert all([ gabarito[i] == resultado[i] for i in range(4)  ]) # Feedback: Algum dos elementos da lista de teste teve resultado diferente. Confira seu código.

```

# Exercício 3-A
Tuplas e listas são muito populares para elencar elementos. Apesar disso, dicionários são muito versáteis para entradas de bases de dados e geralmente são combinados com listas.

Neste exercício, implemente uma função chamada `cria_base_de_dados` (não precisa colocar nenhum parâmetro) que cria e retorna um dicionário com três chaves do tipo string, `"nome"`, `"P1"` e `"P2"`. Cada chave dessas será inicializada com uma lista vazia. 

**Lembrete:** Existem duas formas de criar um dicionário. Em uma delas, especificamos no momento da criação quais serão as chaves e qual será o valor de cada uma. A outra forma consiste em inicializar o dicionário vazio e acrescentar as chaves depois.
```python
# Primeira forma
dicionario1 = {
    "chave" : valor,
    "outra_chave" : outro_valor,
    ...
    "mais_outra_chave" : mais_um_valor
}

# Segunda forma
dicionario2 = {}
dicionario2["chave"] = valor
dicionario2["outra_chave"] = outro_valor
dicionario2["mais_outra_chave"] = mais_um_valor
```
Perceba que os valores atribuídos podem ser de qualquer tipo, incluindo uma lista. A chave não precisa ser uma `string`, mas para simplificar nesta disciplina, usaremos somente chaves do tipo `string` e `int`.

<!-- grade_id: cell-0f18c273ff79a048, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Use esta célula para testar sua função
print(cria_base_de_dados())
```

<!-- grade_id: cell-8ee85cfe9906b0ff, points: 1, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
_resultado = cria_base_de_dados()

assert "nome" in _resultado # Feedback: A chave "nome" não está contida no seu dicionário
assert "P1" in _resultado # Feedback: A chave "P1" não está contida no seu dicionário
assert "P2" in _resultado # Feedback: A chave "P2" não está contida no seu dicionário

assert type(_resultado["nome"]) is list # Feedback: não foi atribuída uma lista à chave "nome"
assert type(_resultado["P1"]) is list # Feedback: não foi atribuída uma lista à chave "P1"
assert type(_resultado["P2"]) is list # Feedback: não foi atribuída uma lista à chave "P2"

assert len(_resultado["nome"]) == 0 # Feedback: a lista da chave "nome" não está vazia.
assert len(_resultado["P1"]) == 0 # Feedback: a lista da chave "P1" não está vazia.
assert len(_resultado["P2"]) == 0 # Feedback: a lista da chave "P2" não está vazia.
```

# Exercício 3-B
A função que definimos cria uma base de dados de alunos, ou seja, podemos ter acesso a alunos e suas respectivas notas de provas. Apesar disso, pode ser um pouco chato inserir aluno por aluno na base sem um controle pois precisamos garantir que os campos `nome`, `P1` e `P2` sejam inseridos simultaneamente.

Desta forma, **crie o procedimento `insere_aluno` que recebe o dicionário da base de dados, o nome do aluno, a nota da P1 e a nota da P2 e os insere em suas respectivas listas.**

Perceba que não é necessário retornar nenhum valor, visto que as alterações acontecerão direto no dicionário (ele é de tipo mutável).

**Lembrete:** Para acessar uma chave do dicionário use: `dicionario["chave"]`. Se o elemento armazenado nesta posição for uma lista, você pode alterar ela normalmente com `append`, por exemplo: `dicionario["chave"].append(valor)`

<!-- grade_id: cell-1c618b69d5b21f65, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Use esta célular para criar sua base de teste. 
# Caso queira reiniciar sua base de dados, basta rodar novamente esta célula
base_de_dados_teste = cria_base_de_dados()
```

```python
# Use esta célula para testar sua função. Altere os valores a cada vez que rodar esta célula
nome_entrada = "Jacques Pierre"
nota_P1 = 9.5
nota_P2 = 10.0
insere_aluno(base_de_dados_teste, nome_entrada, nota_P1, nota_P2)
print(base_de_dados_teste)
```

<!-- grade_id: cell-e8dae6d6f00902d9, points: 1, grade: true, solution: false -->
```python
_bd_teste = cria_base_de_dados()
insere_aluno(_bd_teste, "Jacques Pierre", 10.0, 9.5)
insere_aluno(_bd_teste, "Albert Norvig", 8.0, 7.5)
insere_aluno(_bd_teste, "Thomas Liebtraum", 6.0, 8.75)

def _isclose(v1,v2,tol=1e-8):
    if abs(v1-v2) < tol:
        return True
    return False

assert _bd_teste["nome"][0] == "Jacques Pierre" # O nome do primeiro aluno não foi inserido corretamente. Verifique se a chave 'nome' na função de inserção está correta.
assert _bd_teste["nome"][1] == "Albert Norvig" # O nome do segundo aluno não foi inserido corretamente. Verifique se a chave 'nome' na função de inserção está correta.
assert _bd_teste["nome"][2] == "Thomas Liebtraum" # O nome do terceiro aluno não foi inserido corretamente. Verifique se a chave 'nome' na função de inserção está correta.

assert _isclose(_bd_teste["P1"][0], 10.0) # A P1 do primeiro aluno não foi inserida corretamente. Verifique se a chave 'P1' na função de inserção está correta.
assert _isclose(_bd_teste["P1"][1], 8.0) # A P1 do segundo aluno não foi inserida corretamente. Verifique se a chave 'P1' na função de inserção está correta.
assert _isclose(_bd_teste["P1"][2], 6.0) # A P1 do terceiro aluno não foi inserida corretamente. Verifique se a chave 'P1' na função de inserção está correta.

assert _isclose(_bd_teste["P2"][0], 9.5) # A P2 do primeiro aluno não foi inserida corretamente. Verifique se a chave 'P2' na função de inserção está correta.
assert _isclose(_bd_teste["P2"][1], 7.5) # A P2 do segundo aluno não foi inserida corretamente. Verifique se a chave 'P2' na função de inserção está correta.
assert _isclose(_bd_teste["P2"][2], 8.75) # A P2 do terceiro aluno não foi inserida corretamente. Verifique se a chave 'P2' na função de inserção está correta.
```

# Exercício 4
Listas em Python vêm com um conteito interessante, o de `slicing`. O `slicing` é interessante para manipulação de listas e elementos e pode ser usado em estruturas de dados mais complexas, como veremos adiante no curso.

Um `slice` é uma "fatia" da lista selecionada por um padrão que nós definimos. Uma forma de acessar os elementos de uma lista por slicing é usando
```python
parte_selecionada = lista[inicio:fim:passo]
```

A variável `parte_selecionada` é uma lista com os elementos de saída do slice da lista `lista`. O valor em `inicio` corresponde ao índice do primeiro elemento a ser selecionado (por padrão é *0*), `fim` é o índice do último elemento a ser selecionado (excluindo ele) - por padrão é o índice do último elemento da lista somado de 1 - e `passo` corresponde ao tamanho do passo entre o elemento atual e o próximo (por padrão é *1*).

Como `inicio`, `fim` e `passo` possuem valores-padrão, podemos omitir eles, gerando outras variações desse tipo de slicing:
```python
parte_selecionada = lista[inicio::] # Começa em "inicio" e termina no último elemento da lista
parte_selecionada = lista[:fim:] # Começa em 0 e termina em "fim" - 1
parte_selecionada = lista[::passo] # Começa em 0 e vai até o último elemento da lista com tamanho de passo "passo"
```

```python
# Use esta célula para testar o funcionamento do slicing.
# Altere os valores no slice para qualquer valor desejado.
lista_slices = [1,2,3,4,5,6,7,8,9]
resultado = lista_slices[5:8:1]
print( resultado )
```

Se `passo` for menor que 0, então os elementos da lista são acessados na ordem contrária. Na verdade, se você tentar acessar uma lista ou tupla com elementos negativos, você irá acessar na ordem contrária.

Neste caso, se for usar os valores de `inicio` e `fim` diferentes do padrão, `fim` precisa ser menor que `inicio`, caso contrário não existirá elementos na lista de saída.

```python
# Use esta célula para testar o funcionamento do slicing.
# Altere os valores no slice para qualquer valor desejado.
lista_slices = [1,2,3,4,5,6,7,8,9]
indice = -1 # Experimente alterar o índice para valores negativos
print(f"Elemento {indice} da lista:", lista_slices[indice])

resultado = lista_slices[8:5:-1] # Experimente alterar os valores de início e fim
print("Resultado slice: ", resultado )
```

Sabendo disto, **implemente o slicing para inverter a ordem dos elementos da lista** `lista_exercicio_4`. Armazene o resultado em `exercicio_4`.

**Dica:** utilize o fato de que com o tamanho de passo menor do que zero os elementos da lista são acessados em ordem contrária e que se não colocar o valor de início e fim os elementos são acessados do último para o primeiro.

```python
# Rode esta célula antes!
lista_exercicio_4 = list( range(10,0,-1) )
print("Lista original: ",lista_exercicio_4)
```

<!-- grade_id: cell-12cf79e4f139c73c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9761f856b3e41aea, points: 1, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert len(exercicio_4) == len(lista_exercicio_4)

gabarito = [1,2,3,4,5,6,7,8,9,10]
assert all([ gabarito[i] == exercicio_4[i] for i in range(len(gabarito))  ]) # Feedback: Algum dos elementos da lista de teste teve resultado diferente. Confira seu código.
```

# Exercício 5 - A
A habilidade de manipular listas e dicionários é importante, mas no cotidiano usamos estruturas de dados mais complexas e o uso de listas e dicionários é mantido para tarefas mais genéricas.

Em áreas da inteligência artificial, como o aprendizado de máquina e visão computacional, precisamos lidar com matrizes e vetores. Para tanto, utilizamos o módulo `numpy` em Python que implementa a maior parte das ferramentas necessárias para se trabalhar com matrizes e vetores de forma eficiente. Além disso, o `numpy` implementa diversas funções e constantes básicas da matemática, incluindo funções trigonométricas, exponenciais e logaritmos.

Vamos começar com o básico. Na célula abaixo, carregue o módulo `numpy` com alias `np`.

<!-- grade_id: cell-0e7f06ca3a8071d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-99b89839d6455ddd, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
import sys
assert "numpy" in sys.modules
```

# Exercício 5 - B
Para criar uma matriz de `N` linhas e `M` colunas usando o Numpy podemos usar as seguintes funções:
```python
dimensoes = (N,M) # As dimensões da matriz são passadas como uma tupla
matriz_vazia = numpy.empty(dimensoes, dtype=tipo) # Cria uma matriz sem inicializar seus elementos
matriz_zeros = numpy.zeros(dimensoes, dtype=tipo) # Cria uma matriz preenchida com zeros
matriz_elems = numpy.full(dimensoes, valor, dtype=tipo) # Cria uma matriz preenchida com 'valor'
matriz_diag = numpy.diag([1, 2, 3], dtype=tipo ) # Cria uma matriz diagonal com os valores da lista passada
```

Fornecer o tipo de cada elemento da matriz (em `dtype=tipo`) não é obrigatório. Apesar disso, fornecer o tipo é uma forma de garantir que os dados que estamos fornecendo estão de acordo com a lógica do código.

Também é possível inicializar uma matriz a partir de uma lista de listas
```python
lista_de_listas = [
    [5,2,3],
    [-1,5,2],
    [0,1,0]
]
matriz = numpy.array(lista_de_listas, dtype=tipo)
```

Sabendo disto, **crie uma matriz de 3 linhas por 3 colunas igual à matriz abaixo usando listas de listas.** Armazene na variável `matriz` e atribua o tipo inteiro usando `dtype=int`.
$$
\begin{bmatrix}
10 & 1   & 10 \\
1  & 100 & 1  \\
10 & 1   & 10
\end{bmatrix}
$$

<!-- grade_id: cell-e2172d73313b5b96, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41182b2118159600, points: 1, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert (matriz[0,0] == 10) and (matriz[2,0] == 10) and (matriz[0,2] == 10) and (matriz[2,2] == 10) # Os 10s estão na posição errada
assert (matriz[1,0] == 1) and (matriz[0,1] == 1) and (matriz[2,1] == 1) and (matriz[1,2] == 1) # Os 1s estão na posição errada
assert matriz[1,1] == 100 # O 100 está na posição errada.
```

Usar matrizes para solução de problemas implica em manipular e acessar seus elementos. Para acessar um elemento de uma matriz, utilizamos `matriz[linha,coluna]`, em que `linha` e `coluna` correspondem à linha e coluna, respectivamente, que queremos acessar.

Para acessar cada elemento individualmente da matriz, precisamos usar duas estruturas de repetição aninhadas, uma para alterar as linhas e outra para alterar as colunas. Normalmente, usamos `for` para isso em conjunto de `range(inicio,fim,passo)`, que gera um iterável de números inteiros que se iniciam em `inicio` (padrão é 0), terminam em `fim-1` e têm passo `passo` (padrão é 1). 

Se omitirmos o inicio e o passo e substituirmos `fim` pelo número de linhas, temos os números que vão de 0 a `numero de linhas - 1`, e o equivalente para o número de colunas. 

Assim, o resultado do laço para acessar cada elemento da matriz é
```python
for linha in range(numero_linhas):
    for coluna in range(numero_colunas):
        # operação em matriz[linha,coluna]
```

Para se obter o número de linhas e o número de colunas de uma matriz, utilizamos o atributo `matriz.shape`, que retorna uma tupla de dois elementos: o primeiro é a linha e o segundo é a coluna.

Algumas formas de se usar esse atributo:
```python
# Forma 1
numero_linhas = matriz.shape[0]
numero_colunas = matriz.shape[1]

# Forma 2
numero_linhas, numero_colunas = matriz.shape
```
A forma 1 é mais intuitiva para quem tem contexto em linguagens como Java e C++. A forma 2 é mais idiomático do Python: ele separa a tupla de dois elementos e atribui individualmente para cada uma das variáveis em sequência, tal que o elemento 0 é atribuído à variável mais à esquerda e o elemento 1 é atribuído à próxima variável.

Sabendo disto, implemente uma função chamada `multiplica`, que recebe uma matriz e um número. Nesta função:
* Crie uma matriz de mesmo tamanho que a matriz de entrada (use o `shape` para ajudar e `dtype=float`)
* Atribua à cada linha e coluna da matriz de saída o elemento da matriz de entrada na mesma linha e coluna multiplicado pelo número de entrada
* Retorne a matriz de saída

<!-- grade_id: cell-3938a5c8512d8432, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Use esta célula para testar o funcionamento de sua função.
matriz_teste = np.full((3,5), 1., dtype=float) # Cria uma matriz preenchida somente com números 1
print("Matriz original")
print(matriz_teste)

# Altere os valores em numero para avaliar a saída.
numero = 10
print("Matriz multiplicada")
print(multiplica(matriz_teste, numero))
```

<!-- grade_id: cell-e0bd50e115b7fe22, points: 1, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
matriz_teste = np.array([
    [2.,1.],
    [3.,5.]
])
gabarito = np.array([
    [24.,12.],
    [36.,60.]
])

assert np.all( multiplica(matriz_teste,12) == gabarito )

```

# Exercício 5 - C
Manipular elemento a elemento de matrizes, se a operação independe da posição, é chato e aumenta muito o número de linhas do código, além do de variáveis. Além disso, em Python, acessar matrizes elemento a elemento pode ser custoso computacionalmente. Assim, o conceito de **vetorização** nos auxilia em cada um desses problemas.

**Vetorização** consiste em aplicar operações que seriam executadas elemento a elemento de um vetor para todos os elementos de uma vez só em uma única operação.

Por exemplo, considere o exercício anterior. Nele, iteramos na matriz para multiplicar cada um dos elementos. Apesar disso, se quisermos multiplicar todos os elementos de uma matriz de uma vez só em Python usando o Numpy, podemos usar a seguinte operação:
```python
matriz_multiplicada = numero * matriz
```

Note a simplicidade do código acima. Em uma linha, fizemos com que cada elemento da matriz fosse alterado sem nenhum `for` explicito e, por menos intuitivo que possa parecer, em Python é mais eficiente operar por vetorização que usando laços explícitos.

Usando o numpy, podemos fazer qualquer operação aritmética de forma vetorizada em uma matriz. Por exemplo:
```python
matriz + numero # Soma elemento a elemento da matriz com o número
matriz - numero # Subtrai de cada elemento da matriz o número
matriz * numero # Multiplica todos os elementos da matriz pelo número
matriz / numero # Divide todos os elementos da matriz pelo número
matriz ** numero  # Eleva todos os elementos da matriz pelo número
```

Além disso, as funções disponíveis em Numpy são vetorizadas. Por exemplo
```python
numpy.sin(matriz) # Calcula o seno de cada elemento da matriz
numpy.cos(matriz) # Calcula o cosseno de cada elemento da matriz
numpy.log2(matriz) # Calcula o log em base 2 de cada elemento da matriz
numpy.log(matriz) # Calcula o log em base usando constante de Euler de cada elemento da matriz
numpy.exp(matriz) # Eleva a constante de Euler para cada elemento da matriz
```

Considerando o contexto apresentado, implemente uma função (chamada `onda_cosseno`) em Python que, fornecida uma matriz, um número `a` e um número `b` retorne a matriz que resulta de $a \cdot cos(matriz) + b$.

<!-- grade_id: cell-0adf251123601828, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Use esta célula para testar o funcionamento de sua função.
matriz_teste = np.array([[0., np.pi/2],[np.pi, (3*np.pi/2)]], dtype=float) # Cria uma matriz preenchida
print("Matriz original")
print(matriz_teste)

# Altere os valores para avaliar a saída.
a = 1.0
b = 0.0
print("Matriz após operação")
print(onda_cosseno(matriz_teste, a, b))
```

Para o caso em que `a=1` e `b=0`, a matriz de saída deve ser próxima de
$$
\begin{bmatrix}
1 & 0 \\
-1 & 0
\end{bmatrix}
$$

<!-- grade_id: cell-2cfd3b295f614b6d, points: 1, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
matriz_teste = np.array([[0., np.pi/2],[np.pi, (3*np.pi/2)]], dtype=float)
gabarito = np.array([
    [1.,0.],
    [-1.,0.]
])

assert np.all( 
    np.isclose(
        onda_cosseno(matriz_teste,1,0),
        gabarito 
    )
)

```