# Dados do aluno
Preencha seus dados nas variáveis da célula abaixo para continuar.

```python
# Preencha seu nome e código do aluno aqui
nome_aluno   = ""   # Substitua aqui pelo seu nome, exemplo:   nome_aluno = "Edson Crosué das Dores"
codigo_aluno = None # Substitua aqui pelo seu código, exemplo: codigo_aluno = 1234567
```

<!-- grade_id: cell-02452428bf4b1ae3, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert nome_aluno != "" # Você precisa preencher seu nome no formulário
assert codigo_aluno is not None # Você precisa preencher seu código de aluno
```

# Objetivo
Nesta atividade iremos aplicar os conceitos aprendidos sobre o KNN, DWNN e o KMeans para fazer estimativas na base de dados **Real Estate Valuation Data Set** (usada no exercício anterior). 

Teremos dois objetivos nesta base: 
1. Preparar um modelo de regressão para estimar o preço de um imóvel dados seus atributos
2. Encontrar grupos de imóveis que são relacionados de alguma forma.

Para os nossos objetivos, é necessário que você importe os módulos `numpy` (como `np`), `pandas` (como `pd`) e o `sklearn`.

**Importe os módulos na célula abaixo.**

<!-- grade_id: cell-c760a5a3668ae237, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-239a3c231c571e39, points: 0, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
import sys
assert "numpy" in sys.modules
assert "pandas" in sys.modules
assert "sklearn" in sys.modules
```

# Carregar a Base de Dados
Comece carregando a base de dados Real Estate Valuation (REV) Data Set, do arquivo `real_estate_valuation_data_set.xlsx`, no dataframe `rev_df` usando `pandas.read_excel()`. Para ver se os dados foram carregados corretamente, use o método `head`.

Também coloque a coluna `No` como índice da tabela (lembre-se de que não é uma operação **inplace**).

<!-- grade_id: cell-83860da5098522be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4fa6f450c9a86d28, points: 0, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert rev_df.shape[0] == 414
assert rev_df.shape[1] == 7
assert "No" not in rev_df.columns # Verifique se colocou "No" como índice!
```

# Ajuste da Tabela

* Remova as **colunas** (lembre-se de como usar o comando `drop`) do dataframe de entrada `X1 transaction date`, `X5 latitude` e `X6 longitude` e atribua à mesma variável.
* Troque os nomes das colunas restantes da tabela por `Idade`, `Distância da Estação`, `# Lojas de Conveniência` e `TWD por Ping`.

**Dica:** caso esteja em dúvida sobre como usar os comandos, use as instruções de aula e de exercícios sobre manipulação de dados usando Pandas

<!-- grade_id: cell-1f9723c9fc9e0b4e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd64326ec417b260, points: 0, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert "X1 transaction date" not in rev_df.columns # Veja se você colocou o atributo "X1 transaction date" no drop
assert "X5 latitude" not in rev_df.columns # Veja se você colocou o atributo "X5 latitude" no drop
assert "X6 longitude" not in rev_df.columns # Veja se você colocou o atributo "X6 longitude" no drop

assert "Idade" in rev_df.columns # Veja se você renomeou as colunas
assert "Distância da Estação" in rev_df.columns # Veja se você renomeou as colunas
assert "# Lojas de Conveniência" in rev_df.columns # Veja se você renomeou as colunas
assert "TWD por Ping" in rev_df.columns # Veja se você renomeou as colunas
```

# Separar Base de Treino e Teste
Para treinar e validar nossos algoritmos de ML, precisamos ter uma base de dados para treinamento e outra para teste. 

Os dados de treinamento são responsáveis por fazer o modelo de aprendizado de máquina a ajustar seus parâmetros durante um processo de otimização. No caso do KNN e do DWNN, nenhum parâmetro é ajustado durante a fase de treinamento. Por outro lado, KMeans e outros algoritmos fazem ajustes iterativamente em busca de reduzir o erro de estimação.

Caso a base de dados de interesse não tenha sido previamente separada, costumamos separar 75% dos dados para treinamento e 25% dos dados para teste (isso varia de acordo com o tamanho da base de dados).

Para o pandas, usaremos o método `tabela.sample(frac=, random_state=)`, o qual amostra as linhas sem reposição a porcentagem indicada em `frac` (valor entre 0 e 1, sendo que 1 corresponde a 100%). O resultado deve ser armazenado em outra variável (armazene em `rev_df_treino`). Sete o parâmetro `random_state=1` para este exercício.

Com a tabela amostrada, podemos obter a base de teste removendo da tabela original as linhas que foram amostradas usando `drop(sequencia_de_indices)`. Essa sequência de índices pode ser obtida por meio de `tabela_de_amostras.index`.

**Sabendo disto, aplique os conceitos para amostrar 75% dos dados para a base de treino (armazene em `rev_df_treino`) e o restante, para teste, armazene em `rev_df_teste`.**

<!-- grade_id: cell-3d101639b46edc5a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-16c7a9a84959507b, points: 0, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert rev_df_treino.shape[0] + rev_df_teste.shape[0] == rev_df.shape[0]
assert np.isclose( rev_df_treino.shape[0] / rev_df.shape[0], 0.75,atol=0.01)
```

# Exercício 1: Usando o KNN para regressão
Agora que temos nossa base de dados pronta e separada em treino e teste, podemos começar a utilizar nossos modelos de aprendizado de máquina. 

Vamos começar pelo problema de regressão: *Fornecidos atributos como entrada, queremos estimar o preço da área do imóvel.*

Para esta tarefa usaremos a implementação do KNN pelo `sklearn`. Para usarmos os modelos supervisionados de machine learning do `Scikit Learn`, existe um padrão:
* Primeiro instanciamos o modelo (neste caso será `sklearn.neighbors.KNeighborsRegressor()`).
* Em seguida, fornecemos a base de treinamento para o modelo pelo método `modelo.fit(tabela_com_atributos, rotulos)`, em que a tabela com atributos não tem a coluna dos rótulos.
* Por fim, para ativar o modelo treinado em uma entrada nunca vista, usamos `modelo.predict(entrada)`, o qual retorna o valor estimado.

Desta forma, para este exercício faça:
1. Separe os atributos de treinamento da tabela `rev_df_treino` usando `drop()` e armazene em `X_treino`
2. Armazene a coluna `"TWD por Ping"` na variável `Y_treino`
3. Importe o módulo `sklearn.neighbors`
4. Instancie o modelo de KNN usando `sklearn.neighbors.KNeighborsRegressor(n_neighbors=X_treino.shape[0],weights="distance")` e armazene na variável `modelo_knn`
5. "Treine" o modelo usando o método `modelo.fit(X, Y)`, substitua `X` por `X_treino` e `Y` por `Y_treino`
6. (Opcional) Execute `print` na predição usando `modelo.predict([[19.5, 306.5947, 9]])`. Este seria o primeiro exemplo da base de teste e o resultado correto seria 42.2. Entretanto, o resultado que é o resultado obtido é um pouco maior, $47.92$.

<!-- grade_id: cell-6c24ba8598e9975e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6bd60b72083f4382, points: 4, grade: true, solution: false -->
```python
# Célula reservada ao autograder
_testes_publicos = [
    [19.5, 306.5947, 9],
    [13.3, 561.9845, 5],
    [13.3, 561.9845, 5]
]
assert "sklearn.neighbors" in sys.modules
assert np.all( 
    np.isclose( 
        modelo_knn.predict(_testes_publicos),
        np.array([47.91760389, 44.42684629, 44.42684629])
    )
)
```

O resultado estimado pelo modelo é diferente do valor real do imóvel. Isto pode ser devido a vários fatores, principalmente:
* **Os dados não estão normalizados** (principal causa neste exercício)
* Os atributos usados para estimativa são insuficientes para uma boa estimativa
* O modelo usado para regressão não é o melhor para o caso apresentado

*Mais adiante neste notebook, veremos o efeito da normalização de dados z-score.*

Para saber o quão bem nosso modelo faz as suas estimativas, usamos a base de testes para fazer a estimativa e comparamos o resultado do modelo com o resultado esperado.

Uma das formas de comparar o resultado obtido do esperado é usando o **coeficiente de determinação**, também conhecido como `score R2` (em inglês ficaria `R2 score`). Ele é calculado da seguinte forma:
$$
    R_2 = 1 - \frac{D_{Euclidiana}^2(Y, \hat{Y})}{ D_{Euclidiana}^2(Y, Y_{médio}) },
$$
em que $Y$ é o resultado esperado, $\hat{Y}$ é o resultado obtido pelo modelo e $Y_{médio}$ é a média do vetor $Y$.

Para entender esta fórmula, basta perceber que, quando a estimativa, $\hat{Y}$, é igual ao resultado esperado, $Y$, temos que $D_{Euclidiana}^2(Y, \hat{Y}) = 0$ e, portanto, o score $R_2=1$. Por outro lado, se o resultado obtido for sempre o mesmo valor independente da entrada, $R_2 = 0$. Independente do valor, quando menor for, pior é o modelo e o modelo perfeito teria $R_2 = 1$.

**Sabendo disto, faça**:
* Separe os atributos dos rótulos da tabela `rev_df_teste` assim como foi feito no exercício  anterior e armazene os valores em `X_teste` e `Y_teste`
* Para calcular o score, aplique `modelo.score(X,Y)`, tal que `X` seria `X_teste` e `Y` seria `Y_teste`. Este método irá aplicar a predição para todas as entradas em `X` e iria comparar usando o score `R2` o rótulo `Y_teste` com o resultado obtido. Armazene o resultado em `score_r2`.

<!-- grade_id: cell-144aae175261644a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-88ed61cf402dad01, points: 1, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert np.isclose(score_r2,0.658,atol=0.01) # O score obtido não está de acordo com o esperado.
```

# Exercício 2: Usando K-Means para Agrupamento
Por um momento, vamos visualizar novamente as distribuições de preço com relação às variáveis `Idade` `Distância da Estação` e `# Lojas de Conveniência`. 

Rode a célula abaixo para visualizar a dispersão dos dados.

```python
import matplotlib.pyplot as plt
cmap = "jet"
fig, axs = plt.subplots(1,3,figsize=(18,4))
rev_df.plot.scatter(y="Idade", x="Distância da Estação",c="TWD por Ping", colormap=cmap, ax=axs[0])
rev_df.plot.scatter(y="Idade", x="# Lojas de Conveniência",c="TWD por Ping", colormap=cmap, ax=axs[1])
rev_df.plot.scatter(y="Distância da Estação", x="# Lojas de Conveniência",c="TWD por Ping", colormap=cmap, ax=axs[2])
plt.show()
```

Fora o gráfico mais à esquerda, é um pouco difícil de - visualmente - encontrar padrões de agrupamento no gráfico. Neste ponto, o KMeans é uma ferramenta muito útil.

Para incluir o KMeans em seu projeto, você precisa importar o módulo `sklearn.cluster`. Faça isso na célula abaixo.

<!-- grade_id: cell-1a01f1ca98945bd2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bc64e56a576e1b16, points: 0, grade: true, solution: false -->
```python
# Célula reservada ao autograder
assert "sklearn.cluster" in sys.modules # É necessário importar o módulo "sklearn.cluster" antes de prosseguir.
```

Para usar o KMeans, não será necessário separar a base em treinamento e teste, visto que não temos grupos pré-definidos.

Desta forma
1. **Instancie o modelo usando `sklearn.cluster.KMeans()`** com argumentos `n_clusters=4` e `random_state=1` e armazene em `modelo_kmeans`
2. **Adapte o modelo aos dados usando `modelo.fit(dados)`**, com `dados` sendo a tabela `rev_df`.

<!-- grade_id: cell-a6e7ed678149db7a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b0179e114f6357ce, points: 1, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
clusters = np.array([
 [1.81932836e+01,3.59555818e+02,5.48134328e+00,4.49339552e+01],
 [1.29547619e+01,2.29370698e+03,2.42857143e+00,2.70714286e+01],
 [1.75805970e+01,1.37764252e+03,1.71641791e+00,2.76507463e+01],
 [1.98702703e+01,4.42513443e+03,2.43243243e-01,1.87000000e+01]
])
assert np.all(
    np.isclose(clusters, modelo_kmeans.cluster_centers_)
)
```

A partir desta célula seu modelo está treinado, ou seja, fornecido um elemento podemos estimar a qual grupo ele pertence. 

Vamos verificar a qual grupo pertence cada objeto da base de dados: aplique `modelo.predict(tabela)` e armazene na variável `grupos`. Rode `print` na variável `grupos` para visualizar o que é o resultado do KMeans.

<!-- grade_id: cell-dd136f52bf743d8c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd3f241efc40d0da, points: 2, grade: true, solution: false -->
```python
# Célula reservada ao autograder
assert grupos[0] == 0
assert grupos[-1] == 0
assert grupos[5] == 1

```

O resultado obtido deve ser uma lista contendo números inteiros entre 0 e 3. O que esta lista indica é que:
* Os elemento nas posições de 0 a 4 (linhas 0, 1, 2, 3 e 4 da tabela) pertencem ao grupo 0
* O elemento na posição 5 pertence ao grupo 1
* O mesmo vale para os demais elementos da lista

Mas afinal, o que são esses grupos? Vamos visualizar no gráfico abaixo.

```python
cmap="plasma"
rev_df_viz = rev_df.copy()
rev_df_viz["Grupo"] = rev_df.apply(lambda row: modelo_kmeans.predict([row])[0],axis="columns")
fig, axs = plt.subplots(2,3,figsize=(18,8))
rev_df_viz.plot.scatter(y="Idade", x="Distância da Estação",c="Grupo", colormap=cmap, ax=axs[0,0])
rev_df_viz.plot.scatter(y="Idade", x="# Lojas de Conveniência",c="Grupo", colormap=cmap, ax=axs[0,1])
rev_df_viz.plot.scatter(y="Idade", x="TWD por Ping",c="Grupo", colormap=cmap, ax=axs[0,2])
rev_df_viz.plot.scatter(y="Distância da Estação", x="# Lojas de Conveniência",c="Grupo", colormap=cmap, ax=axs[1,0])
rev_df_viz.plot.scatter(y="Distância da Estação", x="TWD por Ping",c="Grupo", colormap=cmap, ax=axs[1,1])
rev_df_viz.plot.scatter(x="# Lojas de Conveniência", y="TWD por Ping",c="Grupo", colormap=cmap, ax=axs[1,2])
plt.show()
```

Os grupos exibidos não parecem fazer muito sentido. Isto se deve ao fato de os dados terem ordem de magnitude diferentes.

# Exercício 3: Normalização Z-Score
Perceba a diferença da ordem de magnitude dos atributos:
* Os valores em `Idade` não ultrapassam 50
* Em `# Lojas de Conveniência` os valores variam de 0 a 10
* Para `TWD por Ping` os valores estão entre 0 e 120
* Para `Distância`, os valores começam em 0 e quase alcançam 7000.

Nós temos que o KMeans é um algoritmo que se baseia na diferença de distância para adaptar seus clusters, logo as mudanças nos atributos que não são a distância vão ter pouco impacto no ajuste dos centros.

Existem formas de contornar este problema. Uma delas é a normalização `z-score`, que normaliza cada atributo para que a distribuição tenha média 0 e desvio padrão 1.

Ela é calculada da seguinte forma para cada coluna: 
* Calcula-se a média $\mu$ da coluna e o desvio padrão $\sigma$
* Para cada exemplo da coluna, calcula-se $x_{novo} = \frac{x - \mu}{\sigma}$

Felizmente, o `scipy.stats` implementa a normalização por nós por meio da função `scipy.stats.zscore(tabela)`, a qual retorna a tabela com cada atributo normalizado.

**Sabendo disto**: 
* Importe o módulo `scipy.stats`
* Aplique a normalização z-score na tabela `rev_df` e 
* Atribua a tabela retornada à variável `rev_df_norm`

<!-- grade_id: cell-c961fffeffdf9509, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91b9463d28061099, points: 1, grade: true, solution: false -->
```python
# Célula reservada ao autograder
assert np.all( np.isclose(rev_df_norm.mean(), 0.0) )
assert np.all( np.isclose(rev_df_norm.std(), 1.0, atol=0.1) )
```

**Instancie novamente o KMeans (atribua à variável `modelo_kmeans_norm` - não esqueça que `n_clusters=4` e `random_state=1`) e treine usando a tabela normalizada, `rev_df_norm`**

<!-- grade_id: cell-0e96418c2e528208, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-039f9b44fcab92bc, points: 1, grade: true, solution: false -->
```python
# Célula reservada ao autograder
clusters = np.array([
    [-0.22507907,  0.26121706, -0.69951375, -0.64106318],
    [ 1.34623608, -0.53486823,  0.62059678,  0.21256786],
    [ 0.17685997,  2.57058916, -1.31515975, -1.36020105],
    [-0.81350805, -0.62782866,  0.64029125,  0.89424981]
])
assert np.all(
    np.isclose(
        modelo_kmeans_norm.cluster_centers_, clusters
    )
)
```

Vamos visualizar como ficaram os agrupamentos após o treinamento com a tabela normalizada

```python
# Cria uma tabela para visualização - não alterar o df original!
rev_df_viz = rev_df.copy()
# Gera a coluna de grupos a partir da tabela normalizada e da aplicação do KMeans nela
rev_df_viz["Grupo"] = rev_df_norm.apply(lambda row: modelo_kmeans_norm.predict([row])[0],axis="columns")

cmap="plasma"
fig, axs = plt.subplots(2,3,figsize=(18,8))
rev_df_viz.plot.scatter(y="Idade", x="Distância da Estação",c="Grupo", colormap=cmap, ax=axs[0,0])
rev_df_viz.plot.scatter(y="Idade", x="# Lojas de Conveniência",c="Grupo", colormap=cmap, ax=axs[0,1])
rev_df_viz.plot.scatter(y="Idade", x="TWD por Ping",c="Grupo", colormap=cmap, ax=axs[0,2])
rev_df_viz.plot.scatter(y="Distância da Estação", x="# Lojas de Conveniência",c="Grupo", colormap=cmap, ax=axs[1,0])
rev_df_viz.plot.scatter(y="Distância da Estação", x="TWD por Ping",c="Grupo", colormap=cmap, ax=axs[1,1])
rev_df_viz.plot.scatter(x="# Lojas de Conveniência", y="TWD por Ping",c="Grupo", colormap=cmap, ax=axs[1,2])
plt.show()
```

# Análise para conclusão
Agora sim a visualização faz mais sentido. Vamos tentar verificar quais tipos de imóveis existem, considerando os grupos gerados:
* Os grupos 1 e 3 (roxo e amarelo) correspondem aos imóveis de valor maior. Esses imóveis possuem um número maior de lojas de conveniência nos arredores e estão localizados mais próximos à estação de trem mais próxima. Ambos divergem somente na idade: o grupo roxo (1) corresponde aos imóveis mais antigos (de 20 ou mais), enquanto o amarelo (3) corresponde aos mais novos (de 20 anos ou menos).
* Os grupos 0 e 2 (azul e laranja) correspondem a imóveis de valor menor, localizados mais distantes da estação mais próxima. Os imóveis do grupo 2 (laranja) são os imóveis mais desvalorizados. Não parece haver distinção entre ambos os grupos por idade, mas parece que existem mais lojas de conveniência próximos aos do grupo 0 do que com relação ao grupo 2.

Por fim, a análise varia de acordo com a experiência do analista e do contexto.

```python

```