# Dados do aluno
Preencha seus dados nas variáveis da célula abaixo para continuar.

```python
# Preencha seu nome e código do aluno aqui
nome_aluno   = ""   # Substitua aqui pelo seu nome, exemplo:   nome_aluno = "Edson Crosué das Dores"
codigo_aluno = None # Substitua aqui pelo seu código, exemplo: codigo_aluno = 1234567
```

<!-- grade_id: cell-4bf0677fb3533055, points: 0, grade: true, solution: false -->
```python
# Não alterar: célula reservada ao autograder
assert nome_aluno != "" # Você precisa preencher seu nome no formulário
assert codigo_aluno is not None # Você precisa preencher seu código de aluno
```

Neste notebook, iremos trabalhar com a biblioteca `pandas` e o `numpy`. Desta forma, carregue os dois módulos na célula abaixo, e use os **alias**: `pd` para o `pandas` e `np` para o `numpy`. 

**Nomear os aliases é importante neste notebook para usar o autograder.**

<!-- grade_id: cell-55538f0248630b4a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d922fd547f0d04ed, points: 0, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
import sys
assert "numpy" in sys.modules
assert "pandas" in sys.modules
```

# Sobre
Este notebook visa guiar o seu estudo de manipulação de dados em Python usando Numpy e Pandas.

Base de dados usada para os exercícios: [Real estate valuation data set](https://archive.ics.uci.edu/ml/datasets/Real+estate+valuation+data+set).

## Real Estate Valuation Data Set
Nesta base de dados, o intuito é estimar o impacto das variáveis nos preços de imóveis (regressão). Apesar disso, neste notebook, iremos somente manipular e analisar os dados.

As variáveis contidas neste dataset

| Variável | Atributo                               | Unidade                     | Exemplo                |
| -------- | -------------------------------------- | --------------------------- | ---------------------- |
| No       | O ídentificador único da casa          | Sem unidade                 | 4
| X1       | A data de operação                     | Ano.Proporção de Meses      | 2013.250=Março de 2013 |
| X2       | Idade da casa                          | Anos                        | 32                     |
| X3       | A distância até a estação mais próxima | Metros                      | 84.87882               |
| X4       | Número de lojas de conveniência nos arredores | Inteiro              | 10                     | 
| X5       | Latitude                               | Graus                       | 24.98298               |
| X6       | Longitude                              | Graus                       | 121.54024              |
| Y        | Preço por unidade de área local        | Novo dolar do Taiwan / Ping | 37.9                   |

* Ping é uma medida local que corresponde a 3.3 metros quadrados.

# Exercício 1
Para podermos analisar os dados, precisamos carregá-los. O **Real Estate Valuation Data Set** está em formato `xlsx` mas, apesar disso, o `pandas` é capaz de carregá-lo por meio da instrução `pandas.read_excel(caminho_arquivo)`.

Na célula abaixo, **carregue o arquivo `real_estate_valuation_data_set.xlsx` e armazene o `dataframe` carregado na variável `df_exercicio_1`**.

* Caso seja apontado um `ImportError` durante a operação, instale o `xlrd` e o `openpyxel` pelo Anaconda usando `conda install -c conda-forge xlrd openpyxel`

<!-- grade_id: cell-05769eb822e9d059, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# IMPLEMENTE SUA SOLUÇÃO NAS LINHAS ACIMA
# Exibe as 3 primeiras linhas da tabela
df_exercicio_1.head(3)
```

<!-- grade_id: cell-2ed7bcb34dff56b1, points: 2, grade: true, solution: false -->
```python
# Célula reservada ao autograder!

assert df_exercicio_1 is not None # Feedback: você provavelmente não carregou a base de dados na variável 'df_exercicio_1'
assert "X1 transaction date" in df_exercicio_1.columns
assert "X2 house age" in df_exercicio_1.columns
assert "X3 distance to the nearest MRT station" in df_exercicio_1.columns
assert "X4 number of convenience stores" in df_exercicio_1.columns
assert "X5 latitude" in df_exercicio_1.columns
assert "X6 longitude" in df_exercicio_1.columns
assert "Y house price of unit area" in df_exercicio_1.columns
assert np.all( np.array(df_exercicio_1.shape) == np.array((414,8)) )
assert np.isclose( df_exercicio_1["Y house price of unit area"].min() , 7.6 )
assert np.isclose( df_exercicio_1["Y house price of unit area"].max() , 117.5 )

```

Se tudo ocorreu corretamente na célula acima, uma tabela com três linhas e 8 colunas foi exibida.

# Exercício 2

**Seu objetivo agora é melhorar um pouco a tabela para análise de dados.**

Passos: 
1. As colunas `No`, `X1 transaction date`, `X5 latitude` e `X6 longitude` não serão de interesse para esse exercício. Remova elas usando o método `drop(lista_colunas, axis="column")`.
2. Alterar os nomes das colunas, de forma que 
    * `X2 house age` seja `Idade`,
    * `X3 distance to the nearest MRT station` seja `Distância Estação`, 
    * `X4 number of convenience stores` seja `# Lojas Conveniência`, 
    * `Y house price of unit area` seja `Preço`

### Recomendação de Passos

Primeiro use o método `drop` do objeto `df_exercicio_1` usando a lista `["No", "X1 transaction date", "X5 latitude", "X6 longitude"]` (lembre-se de setar `axis="columns"`) e armazene o resultado em `df_exercicio_2`

Em seguida, atribua ao membro `columns` do objeto `df_exercicio_2` a lista `["Idade", "Distância Estação", "# Lojas Conveniência", "Preço"]`.

<!-- grade_id: cell-ffd63f72827f1758, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# IMPLEMENTE SUA SOLUÇÃO NAS LINHAS ACIMA
# Visualização da tabela
df_exercicio_2.head(3)
```

<!-- grade_id: cell-d1ed4e13a545e913, points: 2, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert df_exercicio_2 is not None
assert "Idade" in df_exercicio_2.columns
assert "Distância Estação" in df_exercicio_2.columns
assert "# Lojas Conveniência" in df_exercicio_2.columns
assert "Preço" in df_exercicio_2.columns

assert "X1 transaction date" not in df_exercicio_2.columns
assert "X2 house age" not in df_exercicio_2.columns
assert "X3 distance to the nearest MRT station" not in df_exercicio_2.columns
assert "X4 number of convenience stores" not in df_exercicio_2.columns
assert "X5 latitude" not in df_exercicio_2.columns
assert "X6 longitude" not in df_exercicio_2.columns
assert "Y house price of unit area" not in df_exercicio_2.columns

```

# Exercício 3
No Brasil, não estamos acostumados a trabalhar com a medida de área "ping" e nem com o valor de dolar do Taiwan. 

Como temos costume de trabalhar com área em metros quadrados e como usamos o real no Brasil, vamos padronizar a medida de preço atual, que está em $TWD/ping$ para $BRL/m^2$.

**Desta forma, altere a coluna de preço de acordo.**

Passos:
1. Crie uma cópia do dataframe `df_exercicio_2` e armazene no dataframe `df_exercicio_3` usando o método `tabela.copy()`.
2. Para a coluna `Preço` do `df_exercicio_2`, aplique a relação: $Preço_{novo} = Preço * 0.19 / 3.3$ (1TWD = 0.19BRL e 1 ping = 3.3$m^2$)
    * Considere fazer essa conversão de maneira vetorizada
3. Coloque o resultado da operação acima na coluna "Preço" do dataframe `df_exercicio_3`, ou seja, em `df_exercicio_3["Preço"]`

Os valores de conversão de dólar do Taiwan para real brasileiro usados neste exercício foram obtidos no dia 4 de agosto de 2021.

<!-- grade_id: cell-b4f44d40152b4af4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# IMPLEMENTE SUA SOLUÇÃO NAS LINHAS ACIMA
# Visualização da tabela!
df_exercicio_3.head(3)
```

<!-- grade_id: cell-e492e9c0eaf5d882, points: 2, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert np.isclose( df_exercicio_3["Preço"].loc[0] , 2.182121) # Confira se sua conta está correta!
assert np.isclose( df_exercicio_3["Preço"].loc[2] , 2.723333) # Confira se sua conta está correta!
assert np.all( np.array(df_exercicio_3.shape) == np.array((414,4)) ) # Por um acaso você alterou o tamanho da tabela?
```

# Exercício 4
Agora que o imóvel está em unidades de medida no contexto brasileiro na tabela `df_exercicio_3`, vamos analisar um pouco sobre este preço.

**Neste exercício, ordene a tabela em ordem descrescente de preço e armazene os 5 imóveis com o preço mais alto na variável `df_top_5_maiores_precos`**

**Dicas**
* Utilize o método `tabela.sort_values(by, ascending)`, setando o argumento `by` com o nome da coluna usada para ordenar a tabela (no caso "Preço"). O argumento `ascending` indica se a ordem será crescente (`ascending=True`) ou descrescente(`ascending=False`). **O resultado desta operação não é feito diretamente na tabela!** Atribua o resultado a outra variável.
* Com a tabela ordenada, para recuperar os k primeiros elementos utilize `tabela.iloc[:k]`. O atributo `iloc` indexa as linhas da tabela pela posição delas. Assim, a primeira linha da tabela pode ser acessada por `tabela.iloc[0]`. Podemos acessar mais de uma linha usando slices da mesma forma que fariamos em vetores.

<!-- grade_id: cell-332255ae135124af, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# IMPLEMENTE SUA SOLUÇÃO NAS LINHAS ACIMA
# Visualização da tabela
df_top_5_maiores_precos.head()
```

<!-- grade_id: cell-a374c8bd8acfa34d, points: 2, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert np.isclose(df_top_5_maiores_precos["Preço"].loc[270], 6.765152)
assert np.isclose(df_top_5_maiores_precos["Preço"].loc[220], 4.508182)
assert np.isclose(df_top_5_maiores_precos["Preço"].loc[105], 4.087879)
assert np.all( np.array(df_top_5_maiores_precos.shape) == np.array((5,4)) ) # Você se lembrou de selecionar somente os 5 primeiros?

```

# Exercício 5
Aparentemente o valor dos imóveis no Taiwan é muito baixo. Se os passos anteriores estão corretos, o maior preço do imóvel por metro quadrado está em aproximadamente $R\$6,80$!

E quanto aos imóveis de menor valor?

**De forma semelhante ao exercício 4, ordene a tabela em ordem crescente de preço e armazene os 5 imóveis com o preço mais baixo na variável `df_top_5_menores_precos`**.

<!-- grade_id: cell-e7d9c9e07e53c10d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# IMPLEMENTE SUA SOLUÇÃO NAS LINHAS ACIMA
# Visualização da tabela
df_top_5_menores_precos.head()
```

<!-- grade_id: cell-943236e7cccf53fa, points: 2, grade: true, solution: false -->
```python
# Célula reservada ao autograder!
assert np.isclose(df_top_5_menores_precos["Preço"].loc[113], 0.437576)
assert np.isclose(df_top_5_menores_precos["Preço"].loc[162], 0.667879)
assert np.isclose(df_top_5_menores_precos["Preço"].loc[330], 0.736970)
assert np.all( np.array(df_top_5_menores_precos.shape) == np.array((5,4)) ) # Você se lembrou de selecionar somente os 5 primeiros?
```

# Análise dos dados (Opcional)
Pelo que podemos perceber, em geral a distância à estação mais próxima e o número de lojas de conveniência nos arredores impactam de alguma forma no preço final dos imóveis. Apesar disso, temos algumas excessões visíveis em nossos dados, como - por exemplo - o imóvel de índices 271 e 114.

O imóvel 271 tem o maior valor, o que pode ser impactado pela distância à estação mais próxima, mas não é um imóvel nem tão velho e nem tão novo quanto os demais. Por outro lado o número de lojas de conveniência é baixo (1). Provavelmente sua localização deve ter sido mais impactante no preço, mas como descartamos esse dado para nossa análise.

Por outro lado, o imóvel 114 tem o menor valor, mas é mais próximo à estação de trem relativo aos demais imóveis mais baratos e mais novo, além de ter um número alto de lojas de conveniência nos arredores.

Vamos analisar visualmente como está distribuído o valor dos móveis com relação à cada variável. **Rode a célula abaixo para gerar os gráficos.**

```python
import matplotlib.pyplot as plt
fig, axs = plt.subplots(1,3, figsize=(12, 4))
df_exercicio_3.plot.scatter(x="Idade", y="Preço", ax=axs[0])
df_exercicio_3.plot.scatter(x="Distância Estação", y="Preço",  ax=axs[1])
df_exercicio_3.plot.scatter(x="# Lojas Conveniência", y="Preço",  ax=axs[2])
```

Olhar individualmente cada variável podemos estimar que:
* Existe uma tendência de aumento de preço para imóveis mais próximos de sua estação de trem mais próxima e com mais lojas de conveniência ao redor.
* Com relação ao preço pela idade, os dados parecem formar uma tendência de parábola, de forma que imóveis muito novos ou muito antigos possuem maior valor.

E se explorarmos os pares de variáveis?

```python
fig, axs = plt.subplots(1,3, figsize=(16, 4))
cmap = "jet"
df_exercicio_3.plot.scatter(y="Idade", x="Distância Estação", c="Preço", ax=axs[0], colormap=cmap)
df_exercicio_3.plot.scatter(y="Idade", x="# Lojas Conveniência",  c="Preço", ax=axs[1], colormap=cmap)
df_exercicio_3.plot.scatter(x="# Lojas Conveniência", y="Distância Estação", c="Preço", ax=axs[2], colormap=cmap)
```

As conclusões sobre cada análise vão da experiência de cada analista e da forma de exploração dos dados.

Normalmente, quando queremos entender melhor sobre o impacto das variáveis ou sobre os proprios dados, usamos técnicas de aprendizado de máquina.

```python

```